```python
%matplotlib inline
%precision 16
from __future__ import print_function
import math
import cmath
import numpy
import numpy.fft
import numpy.linalg
import scipy
import scipy.linalg
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Your name and collaborators below:

# HW 1: Review

## Question 0 [5]

Install Anaconda/Python/Jupyter on your computer.

## Question 1 [25]

**(a)** [5] Write a function that takes in a function $f(x)$, end points $a$ and $b$, and an integer $N$, and evaluates the function at equally-spaced points *including* the boundaries $a$ and $b$.  $N$ should represent the number of points used in the interior of the domain, i.e., you should end up with the evaluation with $N+2$ points.  Do not use `linspace`.

<!-- grade_id: 1a-A, points: 0, grade: false, solution: true -->
```python
def eval_linspace(f, a, b, N):
    """Evaluate the function f at equally-spaced points
    
    Parameters
    ----------
    param1 : type
        Description
    param2 : type
        Description
    param3 : type
        Description
    param4 : type
        Description
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    return f_eval
```

<!-- grade_id: 1a-T, points: 5, grade: true, solution: false -->
```python
# Do not delete this:
f = lambda x: numpy.cos(numpy.pi * x / 2.0)
N = 20
numpy.testing.assert_allclose(eval_linspace(f, -1.0, 1.0, N), f(numpy.linspace(-1.0, 1.0, N + 2)))
print("Success!")
```

**(b)** [10] Write another function that computes the second derivative of a given function $f(x)$ using second order finite differences.  You can use the function from part (a) to do this (and are in fact encouraged to do so).  This function should accept a function $f(x)$, the number of grid points to be used, and the end points of the interval $a$ and $b$.  Make sure to also compute the derivative at $x = a$ and $x = b$ using an appropriate finite difference.

<!-- grade_id: 1b-A, points: 0, grade: false, solution: true -->
```python
def compute_dbl_prime(f, a, b, N):
    """Evalute the second derivative of the function f using a 
    second order finite difference approximation.

    Parameters
    ----------
    param1 : type
        Description
    param2 : type
        Description
    param3 : type
        Description
    param4 : type
        Description
    """
    ### BEGIN SOLUTION
    
    ### END SOLUTION
    return f_dbl_prime
```

<!-- grade_id: 1b-Ta, points: 10, grade: true, solution: false -->
```python
# Do not delete this:
sigma = 0.2
f = lambda x: numpy.exp(-x**2 / sigma**2)
f_dbl_prime = lambda x: f(x) * 2.0 * (2.0 * x**2 / sigma**4 - 1.0 / sigma**2)
N = 500
a = -1.0
b = 1.0
x = numpy.linspace(a, b, N + 2)
error = numpy.abs(compute_dbl_prime(f, a, b, N) - f_dbl_prime(x))
print(("||Error|| = %s" % numpy.linalg.norm(error, numpy.infty)))
assert(numpy.linalg.norm(error, numpy.infty) < 1e-2)
print("Success!")
```

**(c)** [10] For the function

$$
    f(x) = 5 + 2x + x^2 + \sin(3x),
$$

compute and plot the error in the domain $\Omega=[a, b]$ with $a=0, b=800$ between your function from part (b) and the true second derivative in the $\ell_\infty$ norm for $\Delta x = 2^{-n}$ where $n = 4, 5, \ldots, 8$.

<!-- grade_id: 1c-A1, points: 8, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

# Question 2 [30]

**(a)** [10] Consider the following matrix $A$,

$$
    A = 
    \begin{pmatrix}
    2 & 1 & 1 & 0 \\
    4 & 3 & 3 & 1 \\
    8 & 7 & 9 & 5 \\
    4 & 7 & 9 & 8
    \end{pmatrix}.
$$

Assemble the matrix $A$ and compute its determinant using the appropriate `numpy.linalg` method, and verifiy that it is equal to 4. Compute the product of the eigenvalues of $A$ and check that you obtain the same value.

<!-- grade_id: 1c-A1, points: 8, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

**(b)** [10] The exponential of a $N\times N$ matrix $M$ is another $N\times N$ matrix defined by

$$
    e^M = \sum_{k=0}^{\infty}\frac{M^k}{k!}.
$$

where $M^0=I$ (the $N\times N$ identity matrix). Write a function that takes in a matrix $M$ and an integer $l$, and evaluates an approximation of $e^M$ using powers of $M$ up to $M^l$.

```python
def compute_matrix_exp(M, l):
    """Evalute the matrix exponential

    Parameters
    ----------
    param1 : type
        Description
    param2 : type
        Description
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
    return matrix_exp
```

<!-- grade_id: 1c-A1, points: 8, grade: true, solution: true -->
```python
# Do not delete this:
A = numpy.array([[1, 2], [3, 4]])
error = compute_matrix_exp(A, 1000) - scipy.linalg.expm(A)
assert(numpy.linalg.norm(error, numpy.infty) < 1e-6)
print("Success!")
```

**(c)** [10] For any integer $N\geq 0$ let $\omega=e^{2\pi i/N}$ denote a $N$-th root of unity ($\omega^N=1$). 
Write a function that takes in $N$ and construct the following $N\times N$ matrix,

$$
    W = 
    \begin{pmatrix}
    1&1&1&1&\cdots &1\\
    1&\omega &\omega ^{2}&\omega ^{3}&\cdots &\omega ^{N-1}\\
    1&\omega ^{2}&\omega ^{4}&\omega ^{6}&\cdots &\omega ^{2(N-1)}\\
    1&\omega ^{3}&\omega ^{6}&\omega ^{9}&\cdots &\omega ^{3(N-1)}\\
    \vdots &\vdots &\vdots &\vdots &\ddots &\vdots \\
    1&\omega ^{N-1}&\omega ^{2(N-1)}&\omega ^{3(N-1)}&\cdots &\omega ^{(N-1)(N-1)}
    \end{pmatrix}.
$$

<!-- grade_id: 1c-A1, points: 8, grade: true, solution: true -->
```python
def compute_fft_matrix(N):
    """Compute the FFT matrix of size N

    Parameters
    ----------
    param1 : type
        Description
    """
    ### BEGIN SOLUTION

    ### END SOLUTION
    return fft_matrix
```

<!-- grade_id: 1c-A1, points: 8, grade: true, solution: true -->
```python
# Do not delete this:
N = 20
error = compute_fft_matrix(N) - numpy.fft.fft(numpy.eye(N))
assert(numpy.linalg.norm(error, numpy.infty) < 1e-10)
print("Success!")
```

# Question 3 [30]
Solve the following PDEs showing the relevant steps and plot a few time snapshots of the solution.

**(a)** [15] $u_t = \kappa u_{xx}$ with $x \in \Omega = (0, L)$, the boundary conditions $u_x(0, t) = 0$ and $u_x(L, t) = 0$, and the initial condition $u(x, 0) = 3 + \cos\left(\frac{\pi x}{L} \right) - 2 \cos\left(\frac{5 \pi x}{L} \right)$.  When plotting the solution pick a value for $L$ and play with values of $\kappa$.  Make sure your solution makes sense as $t \rightarrow \infty$.

<!-- grade_id: 3a-A2, points: 5, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

**(b)** [15] Solve the PDE $u_t + \left(\frac{1}{2}u^2\right)_x = 0$ where $x \in \Omega = \mathbb R$ and the initial conditions
$$ 
    u(x,0) = \left \{ \begin{aligned}
        2, && x \leq 0, \\
        5, && x > 0.
    \end{aligned} \right.
$$
In addition to plotting the solution also mark the edges of interesting regions such as the rarefaction fan or the shock.

<!-- grade_id: 3b-A2, points: 5, grade: true, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```